Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All) to avoid typical problems with Jupyter notebooks. **Unfortunately, this does not work with Chrome right now, you will also need to reload the tab in Chrome afterwards**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Please put your name here:

```python
NAME = "Aymane Hachcham"
```

---

## Text Generation with Markov Chains

In this exercise, we want to make a Donald Trump fake tweet generator using a simple Markov chain.

```python
# Load the tweet data:
import gzip, re, sys
tweets = [[sys.intern(y) for y in re.split(r"\s", x.strip())] for x in gzip.open("/data/tweets_realDonaldTrump_sanitized.txt.gz", "rt")]
```

```python
# Explore the data:

tweets[:10]
```

Output:
```
[['Thank', 'you', 'Arkansas!', '#Trump2016', '#SuperTuesday'],
 ['Thank', 'you', 'Virginia!', '#Trump2016', '#SuperTuesday'],
 ['Thank', 'you', 'Alabama!', '#Trump2016', '#SuperTuesday'],
 ['Thank', 'you', 'Tennessee!', '#Trump2016', '#SuperTuesday'],
 ['Thank', 'you', 'Massachusetts!', '#Trump2016', '#SuperTuesday'],
 ['Thank', 'you', 'Georgia!', '#SuperTuesday', '#Trump2016'],
 ['Wow!',
  'Thank',
  'you',
  'Louisville,',
  'Kentucky!',
  '',
  '#VoteTrump',
  'on',
  '3/5/2016!',
  'Lets',
  '#MakeAmericaGreatAgain!',
  'http://somelink.com/',
  'http://someotherlink.com/'],
 ['Lets',
  'go',
  'America!',
  'Get',
  'out',
  '&',
  '#VoteTrump!',
  '#Trump2016',
  '#MakeAmericaGreatAgain!',
  '#SuperTuesday',
  'http://somelink.com/',
  'http://someotherlink.com/'],
 ['MAKE', 'AMERICA', 'GREAT', 'AGAIN!'],
 ['Thank', 'you', 'Columbus,', 'Ohio!', 'http://somelink.com/']]
```

## Collect the necessary statistics for the Markov model

We need the term frequencies to predict the next word given the previous 0...order words.

Use `()` (empty tuple) as a start and stop token. Use tuples as keys in your maps.

For the 0th order, this is simply the word frequency!

<!-- grade_id: cell-99d57eea729afd84, points: 0, grade: false, solution: true -->
```python
# Aggregate the data as necessary for Markov model of order 0...order
def aggregate(tweets, order):
    from collections import defaultdict, Counter
    models = []
    # As 0th order, use the first tokens only. 
    tokens = []
    for tweet in tweets:
        tokens.append(tweet[0])
    counter = Counter(tokens)
    
    models.append({(): counter})
    
    for o in range(1, order+1):
        model = {} # use tuple() -> word (or empty tuple)
        
        if o == 1:
            # Code for the first order:
            for tweet in tweets:
                n_tokens = len(tweet)
            
                for index, key in enumerate(tweet):
                    if n_tokens > (index + 1):
                        word = tweet[index + 1]
                    
                        if key not in model:
                            model[key] = [word]
                        else:
                            model[key].append(word)
                    else:
                        word = ()
                        if key not in model:
                            model[key] = [word]
                        else:
                            model[key].append(word)
            
            for k in model:
                model[k] = dict(Counter(model[k]))
            
            models.append(model)
        
        else:
            # Code for the second order:
            for tweet in tweets:
                n_tokens = len(tweet)
    
                for i, key1 in enumerate(tweet):  
                    if n_tokens > i + 2:
                        key2 = tweet[i + 1]
                        word = tweet[i + 2]
                        if (key1, key2) not in model:
                            model[(key1, key2)] = [word]
                        else:
                            model[(key1, key2)].append(word)
                    
                    elif i + 2 == n_tokens:
                        key2 = tweet[i + 1]
                        word = ()
                        if (key1, key2) not in model:
                            model[(key1, key2)] = [word]
                        else:
                            model[(key1, key2)].append(word)
                            
            for k in model:
                model[k] = dict(Counter(model[k]))
            
            models.append(model)
       
    return models
```

<!-- grade_id: cell-cc42f3c46e849c9d, points: 1, grade: true, solution: false -->
```python
#### AUTOMATIC TESTS
_tmp = aggregate(tweets[:100], order=2)
assert len(_tmp) == 3 and isinstance(_tmp, list), "Wrong result"
assert all(isinstance(x, dict) for x in _tmp), "Wrong result"
assert not () in _tmp[0][()], "0th order must not include the end token"
assert sum(_tmp[0][()].values()) == 100, "0th order incorrect."
assert sum(sum(x.values()) for x in _tmp[1].values()) == sum(len(x) for x in tweets[:100]), "1th order incomplete."
assert sum(sum(x.values()) for x in _tmp[2].values()) == sum(len(x)-1 for x in tweets[:100]), "2nd order incomplete."
```

<!-- grade_id: cell-d12114dd97314f71, points: 1, grade: true, solution: false -->
```python
#### Additional hidden AUTOMATIC TESTS
del _tmp
```

### Train your model

```python
%time model = aggregate(tweets, order=3)
```

Output:
```
CPU times: user 1.17 s, sys: 423 ms, total: 1.59 s
Wall time: 1.59 s
```

```python
# Testin the function how it should be:
model = aggregate(tweets[:2], order=2)

order = len(model) - 1

w1 = random.choice([t for tweet in tweets[:2] for t in tweet])
random_order = random.randrange(order)

# print(max(model[1][w1], key=model[1][w1].get))

f_t = []
for i in range(10):
    w1 = random.choice([t for tweet in tweets[:2] for t in tweet])
    
#     print(max(model[1][w1], key=model[1][w1].get))
    
    if max(model[1][w1], key=model[1][w1].get) == ():
        continue
    else:
        w2 = w1 + ' ' + max(model[1][w1], key=model[1][w1].get) 
    
    f_t.append(w2)

f_t
```

Output:
```
['Virginia! #Trump2016',
 'you Arkansas!',
 'Thank you',
 'Virginia! #Trump2016',
 'you Arkansas!',
 'you Arkansas!',
 'Thank you',
 'Virginia! #Trump2016',
 '#Trump2016 #SuperTuesday']
```

## Make Trump tweet again

Lets make Trump tweet again.

Write a function "trump" that randomly generates trumpesque garbage given the above model, by randomly sampling from the appropriate distribution.

```python
import random
order = 2
# random.choice[i for i in range(order)]

random.randrange(order)

# random.choice(order)
```

Output:
```
1
```

<!-- grade_id: cell-afe8e59e335032ef, points: 0, grade: false, solution: true -->
```python
def trump(model):
    """Generate Trumpesque nonsense from a trained Markov model"""
    import random
    order = len(model) - 1
    output = []
    for i in range(0, 100): # enforce a max length 100, in case your stopping does not work
        
        # YOUR CODE HERE
        random_order = random.randrange(order)
        w1 = random.choice([t for tweet in tweets[:2] for t in tweet])
        if w1 in model[random_order]:
            if max(model[random_order][w1], key=model[random_order][w1].get) == ():
                continue
            else:
                next_word = w1 + ' ' + max(model[random_order][w1], key=model[random_order][w1].get)
        else:
            continue
        
        output.append(next_word)
    return output
```

<!-- grade_id: cell-d7acff8452df539e, points: 1, grade: true, solution: false -->
```python
#### AUTOMATIC TESTS
_tmp = aggregate(tweets[:100], order=2)
_tmp = [trump(_tmp) for x in range(100)]
assert any('http://somelink.com/' in x for x in _tmp), "Does not work right."
assert any('#MakeAmericaGreatAgain' in x for x in _tmp), "Does not work right."
assert any(x in tweets[:100] for x in _tmp), "Some tweet must be reproduced"
assert any(x not in tweets[:100] for x in _tmp), "Some tweet must be fake"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [200], in <cell line: 4>()
      2 _tmp = aggregate(tweets[:100], order=2)
      3 _tmp = [trump(_tmp) for x in range(100)]
----> 4 assert any('http://somelink.com/' in x for x in _tmp), "Does not work right."
      5 assert any('#MakeAmericaGreatAgain' in x for x in _tmp), "Does not work right."
      6 assert any(x in tweets[:100] for x in _tmp), "Some tweet must be reproduced"

AssertionError: Does not work right.
```

<!-- grade_id: cell-8c116fea73afa671, points: 1, grade: true, solution: false -->
```python
#### Additional hidden AUTOMATIC TESTS
```

<!-- grade_id: cell-1f89f2f0ecdced63, points: 1, grade: true, solution: false -->
```python
#### Additional hidden AUTOMATIC TESTS
del _tmp
```

<!-- grade_id: cell-77e0ce6a481cac94, points: 1, grade: true, solution: false -->
```python
#### Additional hidden AUTOMATIC TESTS
```

## Make Donald Trump tweet garbage again

Lets make Donald Trump tweet again. Generate some Trumpesque nonsense:

```python
for i in range(10):
    print(*trump(model))
```

```python

```