Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All) to avoid typical problems with Jupyter notebooks. **Unfortunately, this does not work with Chrome right now, you will also need to reload the tab in Chrome afterwards**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Please put your name here:

```python
NAME = "Aymane Hachcham"
```

---

# Deep Neural Embeddings

In this assignment, we will work with deep neural embeddings (but not train such an embedding, which is much too ressource intensive).

```python
import numpy as np, pandas as pd
import json, gzip
raw = json.load(gzip.open("/data/simpsonswiki.json.gz", "rt", encoding="utf-8"))
titles2, texts, classes, mclasses = [x["title"] for x in raw], [x["text"] for x in raw], [x["c"] for x in raw], [x["m"] for x in raw]
df = pd.read_csv("/data/simpsons-vectors-bert-mean.csv.gz") # BERT mean vectors
titles, labels, vectors = df.values[:,1], df.values[:,2], df.values[:,3:]
del raw, df # free memory
```

```python
import sklearn.preprocessing
vectors = sklearn.preprocessing.normalize(vectors, copy=False)
```

Find the "best" result when clustering with k-means for k=2..20 by ARI.

For reproducibility, use the fixed random seed 0, 1 restarts, and no tolerance.

Log for yourself the time needed to cluster.

Note that on real data, we cannot use ARI.

<!-- grade_id: cell-4d15d59170bcfa75, points: 0, grade: false, solution: true -->
```python
from sklearn.cluster import KMeans
from sklearn.metrics import adjusted_rand_score
best = (0, -1, None, None) # ARI, k, assignment, centers

# YOUR CODE HERE
kmeans_model = KMeans(n_clusters=2.2, random_state=1, tol=0).fit(vectors)
centers = kmeans.cluster_centers_
assignments = kmeans.labels_
```

<!-- grade_id: cell-2c96ab5b69f0d7ad, points: 1, grade: true, solution: false -->
```python

```

# Explore the clustering

Explore the clustering: print each clusters size and the 5 most central documents to each cluster.

<!-- grade_id: cell-00e3f06f36608260, points: 0, grade: false, solution: true -->
```python
def explain(ari, k, assignment, centers):
    print("ARI:", ari)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
explain(*best)
```

<!-- grade_id: cell-b6d9e129a79586d0, points: 1, grade: true, solution: false -->
```python
# Hidden unit tests
```

## Improve the cluster explanation with TF-IDF

Interestingly, TF-IDF is still useful here - what are the important words, now that we only work with 768-dimensional mean vectors? These averaged vectors are not very similar to word vectors anymore (all close together, and close to stop words).

First, get back our old tf-idf data

<!-- grade_id: cell-661efbdc2a3813bd, points: 0, grade: false, solution: true -->
```python
# Prepare Tfidf vectors!
tfidf = None # sparse tf-idf matrix
vocabulary = None # vocabulary
idf = None # IDF values
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-12f3cfde0d170ba4, points: 1, grade: true, solution: false -->
```python
# Automatic tests
assert tfidf.shape[0] == vectors.shape[0]
```

Now write an explain2 function that also prints the most important words for each cluster.

Also use the multi-classifiction information in mclasses to explain the cluster contents in terms of labels.

<!-- grade_id: cell-c2ebdb20b9a71576, points: 0, grade: false, solution: true -->
```python
def explain2(ari, k, assignment, centers, tfidf, idf, vocabulary):
    print("ARI:", ari)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
explain2(*best, tfidf, idf, vocabulary)
```

<!-- grade_id: cell-5feccbef29e5c0da, points: 1, grade: true, solution: false -->
```python
# Hidden unit tests
```

<!-- grade_id: cell-25f04a4397e111a4, points: 1, grade: true, solution: false -->
```python
# Hidden unit tests
```